Fast Evolution of SOS-Independent Multi-Drug Resistance in Bacteria

Author response:

The following is the authors’ response to the original reviews.

Public Reviews:

Review #1:

Summary:

Jin et al. investigated how the bacterial DNA damage (SOS) response and its regulator protein RecA affect the development of drug resistance under short-term exposure to beta-lactam antibiotics. Canonically, the SOS response is triggered by DNA damage, which results in the induction of error-prone DNA repair mechanisms. These error-prone repair pathways can increase mutagenesis in the cell, leading to the evolution of drug resistance. Thus, inhibiting the SOS regulator RecA has been proposed as a means to delay the rise of resistance.

In this paper, the authors deleted the RecA protein from E. coli and exposed this ∆recA strain to selective levels of the beta-lactam antibiotic, ampicillin. After an 8-hour treatment, they washed the antibiotic away and allowed the surviving cells to recover in regular media. They then measured the minimum inhibitory concentration (MIC) of ampicillin against these treated strains. They note that after just 8-hour treatment with ampicillin, the ∆recA had developed higher MICs towards ampicillin, while by contrast, wild-type cells exhibited unchanged MICs. This MIC increase was also observed in subsequent generations of bacteria, suggesting that the phenotype is driven by a genetic change.

The authors then used whole genome sequencing (WGS) to identify mutations that accounted for the resistance phenotype. Within resistant populations, they discovered key mutations in the promoter region of the beta-lactamase gene, ampC; in the penicillin-binding protein PBP3 which is the target of ampicillin; and in the AcrB subunit of the AcrAB-TolC efflux machinery. Importantly, mutations in the efflux machinery can impact the resistance towards other antibiotics, not just beta-lactams. To test this, they repeated the MIC experiments with other classes of antibiotics, including kanamycin, chloramphenicol, and rifampicin. Interestingly, they observed that the ∆recA strains pre-treated with ampicillin showed higher MICs towards all other antibiotics tested. This suggests that the mutations conferring resistance to ampicillin are also increasing resistance to other antibiotics.

The authors then performed an impressive series of genetic, microscopy, and transcriptomic experiments to show that this increase in resistance is not driven by the SOS response, but by independent DNA repair and stress response pathways. Specifically, they show that deletion of the recA reduces the bacterium's ability to process reactive oxygen species (ROS) and repair its DNA. These factors drive the accumulation of mutations that can confer resistance to different classes of antibiotics. The conclusions are reasonably well-supported by the data, but some aspects of the data and the model need to be clarified and extended.

We sincerely appreciate your overall summary of the manuscript and their positive evaluation of our work.

Strengths:

A major strength of the paper is the detailed bacterial genetics and transcriptomics that the authors performed to elucidate the molecular pathways responsible for this increased resistance. They systemically deleted or inactivated genes involved in the SOS response in E. coli. They then subjected these mutants to the same MIC assays as described previously. Surprisingly, none of the other SOS gene deletions resulted in an increase in drug resistance, suggesting that the SOS response is not involved in this phenotype. This led the authors to focus on the localization of DNA PolI, which also participates in DNA damage repair. Using microscopy, they discovered that in the RecA deletion background, PolI co-localizes with the bacterial chromosome at much lower rates than wild-type. This led the authors to conclude that deletion of RecA hinders PolI and DNA repair. Although the authors do not provide a mechanism, this observation is nonetheless valuable for the field and can stimulate further investigations in the future.

In order to understand how RecA deletion affects cellular physiology, the authors performed RNA-seq on ampicillin-treated strains. Crucially, they discovered that in the RecA deletion strain, genes associated with antioxidative activity (cysJ, cysI, cysH, soda, sufD) and Base Excision Repair repair (mutH, mutY, mutM), which repairs oxidized forms of guanine, were all downregulated. The authors conclude that down-regulation of these genes might result in elevated levels of reactive oxygen species in the cells, which in turn, might drive the rise of resistance. Experimentally, they further demonstrated that treating the ∆recA strain with an antioxidant GSH prevents the rise of MICs. These observations will be useful for more detailed mechanistic follow-ups in the future.

We are grateful to you for your positive assessment of the strengths of our manuscript and your recognition of its potential future applications.

Weaknesses:

Throughout the paper, the authors use language suggesting that ampicillin treatment of the ∆recA strain induces higher levels of mutagenesis inside the cells, leading to the rapid rise of resistance mutations. However, as the authors note, the mutants enriched by ampicillin selection can play a role in efflux and can thus change a bacterium's sensitivity to a wide range of antibiotics, in what is known as cross-resistance. The current data is not clear on whether the elevated "mutagenesis" is driven ampicillin selection or by a bona fide increase in mutation rate.

We greatly appreciate you for raising this issue, as it is an important premise that must be clearly stated throughout the entire manuscript. To verify that the observed increase in mutation rate is a bona fide increase and not due to experimental error, we used a non-selective antibiotic, rifampicin, to evaluate the mutation frequency after drug induction, as it is a gold-standard method documented in other studies [Heterogeneity in efflux pump expression predisposes antibiotic-resistant cells to mutation, Science, 362, 6415, 686-690, 2018.]. In the absence of ampicillin treatment, the natural mutation rates detected using rifampicin were consistent between the wild-type and the ΔrecA strain. However, after ampicillin treatment, the mutation rate detected using rifampicin was significantly elevated only in the ΔrecA strain (Fig. 1G). We also employed other antibiotics, such as ciprofloxacin and chloramphenicol, in our experiments to treat the cells (data not shown). However, we observed that beta-lactam antibiotics specifically induced the emergence of resistance or altered the MIC in our bacterial populations. If resistance had pre-existed before antibiotic exposure or a bona fide increase in mutation rate, we would expect other antibiotics to exhibit a similar selective effect, particularly given the potential for cross-resistance to multiple antibiotics.

Furthermore, on a technical level, the authors employed WGS to identify resistance mutations in the treated ampicillin-treated wild-type and ∆recA strains. However, the WGS methodology described in the paper is inconsistent. Notably, wild-type WGS samples were picked from non-selective plates, while ΔrecA WGS isolates were picked from selective plates with 50 μg/mL ampicillin. Such an approach biases the frequency and identity of the mutations seen in the WGS and cannot be used to support the idea that ampicillin treatment induces higher levels of mutagenesis.

We appreciate your concern regarding potential inconsistencies in the WGS methodology. However, we would like to clarify that the primary aim of the WGS experiment was to identify the types of mutations present in the wild-type and ΔrecA strains after treatment of ampicillin, rather than to quantify or compare mutation frequencies. This purpose was explicitly stated in the manuscript.

Furthermore, the choice of selective and non-selective conditions was made to ensure the successful isolation of mutants in both strains. Specifically, if selective conditions (50 μg/mL ampicillin) were applied to the wild-type strain, it would have been nearly impossible to recover colonies for WGS analysis, as wild-type cells are highly susceptible to ampicillin at this concentration (Top, Author response image 1). Conversely, under non-selective conditions, ΔrecA mutants carrying resistance mutations may not have been effectively isolated, which would have limited our ability to identify resistance mutations in these strains (Bottom, Author response image 1 Thus, the use of different selection pressures was essential for achieving the objective of mutation identification in this study.

After 8 hours of antibiotic treatment, the wild type or the ΔrecA cells were plated on agar plates either without ampicillin or with 50 μg/mL ampicillin and incubated for 24-48 hours.

Top: Under selective conditions, no wild type colonies were recovered, indicating high susceptibility to the antibiotic, preventing further analysis. Bottom: In non-selective conditions, both ΔrecA resistant mutants and non-resistant cells grew, making it difficult to distinguish and isolate the mutants carrying resistance mutations.

Finally, it is important to establish what the basal mutation rates of both the WT and ∆recA strains are. Currently, only the ampicillin-treated populations were reported. It is possible that the ∆recA strain has inherently higher mutagenesis than WT, with a larger subpopulation of resistant clones. Thus, ampicillin treatment might not in fact induce higher mutagenesis in ∆recA.

Thanks for this suggestion. The basal mutation frequency of the wild-type and the ∆recA strain have been measured using rifampicin (Fig. 1G), and there is no significant difference between them.

Reviewer #2:

Summary:

This study aims to demonstrate that E. coli can acquire rapid antibiotic resistance mutations in the absence of a DNA damage response. To investigate this, the authors employed a sophisticated experimental framework based on a modified Adaptive Laboratory Evolution (ALE) workflow. This workflow involves numerous steps culminating in the measurement of antibiotic resistance. The study presents evidence that a recA strain develops ampicillin resistance mutations more quickly than the wild-type, as shown by measuring the Minimum Inhibitory Concentration (MIC) and mutation frequency. Whole-genome sequencing of 15 recA-colonies resistant to ampicillin revealed predominantly inactivation of genes involved in the multi-drug efflux pump system, whereas, in the wild-type, mutations appear to enhance the activity of the chromosomal ampC cryptic promoter. By analyzing mutants involved in the SOS response, including a lexA3 mutant incapable of inducing the SOS response, the authors conclude that the rapid evolution of antibiotic resistance occurs in an SOS-independent manner when recA is absent.

Furthermore, RNA sequencing (RNA-seq) of the four experimental conditions suggests that genes related to antioxidative responses drive the swift evolution of antibiotic resistance in the recA-strain.

We greatly appreciate your overall summary of the manuscript and their positive evaluation of our work.

Weaknesses:

However, a potential limitation of this study is the experimental design used to determine the 'rapid' evolution of antibiotic resistance. It may introduce a significant bottleneck in selecting ampicillin-resistant mutants early on. A recA mutant could be more susceptible to ampicillin than the wild-type, and only resistant mutants might survive after 8 hours, potentially leading to their enrichment in subsequent steps. To address this concern, it would be critical to perform a survival analysis at various time points (0h, 2h, 4h, 6h, and 8h) during ampicillin treatment for both recA and wild-type strains, ensuring there is no difference in viability.

We appreciate your suggestion. We measured the survival fraction at 0, 2, 4, 6, and 8 hours after ampicillin treatment. The results show no significant difference in antibiotic sensitivity between the wild-type and ΔrecA strain (Fig. S2). We therefore added a description int the main text, “Meanwhile, after 8 hours of treatment with 50 μg/mL ampicillin, the survival rates of both wild type and ΔrecA strain were consistent (Fig. S2)”.

The observation that promoter mutations are absent in ΔrecA strains could be explained by previous research indicating that amplification of the AmpC genes is a mechanism for E. coli resistance to ampicillin, which does not occur in a recA-deficient background (PMID# 19474201).

We are very grateful to you for providing this reference. We did examine the amplification of the ampC gene in both wild-type and _recA-_deficient strains, but we found no significant changes in its copy number after ampicillin treatment (Author response image 2). Therefore, the results and discussion regarding gene copy number were not included in this manuscript.

Copy number variations of genes in the chromosome before and after exposure to ampicillin at 50 µg/mL for 8 hours in the wild type and ΔrecA strain.

The section describing Figure 3 is poorly articulated, and the conclusions drawn are apparent. The inability of a recA strain to induce the SOS response is well-documented (lines 210 and 278). The data suggest that merely blocking SOS induction is insufficient to cause 'rapid' evolution in their experimental conditions. To investigate whether SOS response can be induced independently of lexA cleavage by recA, alternative experiments, such as those using a sulA-GFP fusion, might be more informative.

Thanks for your suggestion. We agree that detecting the expression level of SulA can provide valuable information to reveal the impact of the SOS system on rapid drug resistance. In addition to fluorescence visualization and quantification of SulA expression, regulating the transcription level of the sulA gene can achieve the same objective. Therefore, in our transcriptome sequencing analysis, we focused on evaluating the transcription level of sulA (Fig. 4E).

In Figure 4E, the lack of increased SulA gene expression in the wild-type strain treated with ampicillin is unexpected, given that SulA is an SOS-regulated gene. The fact that polA (Pol I) is going down should be taken into account in the interpretation of Figures 2D and 2E.

Thank you for your observation regarding the lack of increased SulA gene expression in the wild-type strain treated with ampicillin in Figure 4E. We agree that SulA is typically an SOS-regulated gene, and its expression is expected to increase in response to DNA damage induced by antibiotics like ampicillin. However, in our experimental conditions, the observed lack of increased SulA expression could be due to different factors. One possibility is that the concentration of ampicillin used, or the duration of treatment, was not applicable to induce a strong SOS response in the wild type strain under the specific conditions tested. Additionally, differences in experimental setups such as timing, sampling, or cellular stress responses could account for the lack of a pronounced upregulation of SulA.

You may state that the fact that polA (Pol I) is going down should be taken into account in the interpretation of Figures 3D and 3E, and we agree with you.

The connection between compromised DNA repair, the accumulation of Reactive Oxygen Species (ROS) based on RNA-seq data, and accelerated evolution is merely speculative at this point and not experimentally established.

We greatly appreciate your comments. First, the correlation between DNA mutations and the accumulation of reactive oxygen species (ROS) has been experimentally confirmed. As shown in Fig. 4I, after the addition of the antioxidant GSH, DNA resistance mutations were not detected in the ΔrecA strain treated with ampicillin for 8 hours, compared to those without the addition of GSH, proving that the rapid accumulation of ROS induces the enhancement of DNA resistance mutations. Second, the enhancement of DNA resistance mutations in relation to bacterial resistance has been widely validated and is generally accepted. Finally, we appreciate the your suggestion to strengthen the evidence supporting ROS enhancement. To address this, we have added an experiment to measure ROS levels. Through flow cytometry, we found that ROS levels significantly increased in both the wild-type and ΔrecA strain after 8 hours of ampicillin treatment. However, ROS levels in the ΔrecA strain showed a significant further increase compared to the wild-type strain (Fig. 4G). Additionally, with the addition of 50 mM glutathione, no significant change in ROS levels was observed in either the wild-type or ΔrecA strain before and after ampicillin treatment (Fig. 4H). This result further confirms our finding in Fig. 4I, where adding GSH inhibited the development of antibiotic resistance.

Reviewer #3:

Summary:

In the present work, Zhang et al investigate the involvement of the bacterial DNA damage repair SOS response in the evolution of beta-lactam drug resistance evolution in Escherichia coli. Using a combination of microbiological, bacterial genetics, laboratory evolution, next-generation, and live-cell imaging approaches, the authors propose short-term drug resistance evolution that can take place in RecA-deficient cells in an SOS response-independent manner. They propose the evolvability of drug resistance is alternatively driven by the oxidative stress imposed by the accumulation of reactive oxygen species and inhibition of DNA repair. Overall, this is a nice study that addresses a growing and fundamental global health challenge (antimicrobial resistance). However, although the authors perform several multi-disciplinary experiments, there are several caveats to the authors' proposal that ultimately do not fully support their interpretation that the observed antimicrobial resistance evolution phenotype is due to compromised DNA repair.

We greatly appreciate your overall summary of the manuscript and positive evaluation of our work.

Strengths:

The authors introduce new concepts to antimicrobial resistance evolution mechanisms. They show short-term exposure to beta-lactams can induce durably fixed antimicrobial resistance mutations. They propose this is due to comprised DNA repair and oxidative stress. This is primarily supported by their observations that resistance evolution phenotypes only exist for recA deletion mutants and not other genes in the SOS response.

Thanks for your positive comments.

Weaknesses:

The authors do not show any direct evidence (1) that these phenotypes exist in strains harboring deletions in other DNA repair genes outside of the SOS response, (2) that DNA damage is increased, (3) that reactive oxygen species accumulate, (4) that accelerated resistance evolution can be reversed by anything other than recA complementation. The authors do not directly test alternative hypotheses. The conclusions drawn are therefore premature.

We sincerely thank you for your insightful comments. First, in this study, our primary focus is on the role of recA deficiency in bacterial antibiotic resistance evolution. Therefore, we conducted an in-depth investigation on E. coli strains lacking RecA and found that its absence promotes resistance evolution through mechanisms involving increased ROS accumulation and downregulation of DNA repair pathways. While we acknowledge the importance of other DNA repair genes outside of the SOS response, exploring them is beyond the scope of this paper. However, in a separate unpublished study, we have identified the involvement of another DNA recombination protein, whose role in resistance evolution is not yet fully elucidated, in promoting resistance development. This finding is part of another independent investigation.

Regarding DNA damage and repair, our paper emphasizes that resistance-related mutations in DNA are central to the development of antibiotic resistance. These mutations are a manifestation of DNA damage. To demonstrate this, we measured mutation frequency and performed whole-genome sequencing, both of which confirmed an increase in DNA mutations.

We appreciate the reviewer's suggestion to provide additional evidence for ROS accumulation, and we have now supplemented our manuscript with relevant experiments. Through flow cytometry, we found that ROS levels significantly increased in both the wild type and ΔrecA strains after 8 hours of ampicillin treatment. However, ROS levels in the ΔrecA strain showed a significant further increase compared to the wild-type strain (Fig. 4G). Additionally, with the addition of 50 mM glutathione, no significant change in ROS levels was observed in either the wild-type or ΔrecA strain before and after ampicillin treatment (Fig. 4H). This result further confirms our finding in Fig. 4I, where adding GSH inhibited the development of antibiotic resistance.

Finally, in response to your question about reversing accelerated resistance evolution, we would like to highlight that, in addition to recA complementation, we successfully suppressed rapid resistance evolution by supplementing with an antioxidant, GSH (Fig. 4I). This further supports our hypothesis that increased ROS levels play a key role in driving accelerated resistance evolution in the absence of RecA.

Recommendations for the authors:

Reviewer #1:

The author's model asserts that deletion of recA impairs DNA repair in E. coli, leading to an accumulation of ROS in the cell, and ultimately driving the rapid rise of resistance mutations. However, the experimental evidence does not adequately address whether the resistance mutations are true, de novo mutations that arose due to beta-lactam treatment, or mutations that confer cross-resistance enriched by ampicillin selection.

a. Major: In Figure 1F & G, the authors show that the ∆recA strain, following ampicillin treatment, has higher resistance and mutation frequency towards rifampicin than WT. However, it is not clear whether the elevated resistance and mutagenesis are driven by mutations enriched by the ampicillin treatment (e.g. mutations in acrB, as seen in Figure 2) or by "new" mutations in the rpoB gene. As the authors note, the mutants enriched by ampicillin selection can play a role in efflux and can thus change a bacterium's sensitivity to a wide range of antibiotics, including rifampicin, in what is known as cross-resistance. Therefore, the mutation frequency calculation, which relies on quantifying rifampicin-resistant clones, might be confounded by bacteria with mutations that confer cross-resistance. A better approach to calculate mutation frequency would be to employ an assay that does not require antibiotic selection, such as a lac-reversion assay. This would mitigate the confounding effects of cross-resistance of drug-resistant mutations.

We appreciate your thoughtful comments regarding the potential for cross-resistance to confound the mutation frequency calculation based on rifampicin-resistant clones. Indeed, as noted, ampicillin selection can enrich for mutants with enhanced efflux activity, which may confer cross-resistance to a range of antibiotics, including rifampicin.

However, we believe that the current approach of calculating mutation frequency using rifampicin-resistant mutants is still valid in our specific context. Rifampicin targets the RNA polymerase β subunit, and resistance typically arises from specific mutations in the rpoB gene. These mutations are well-characterized and distinct from those typically associated with efflux-related cross-resistance. Thus, the likelihood of cross-resistance affecting our mutation frequency calculation is minimized in this scenario.

Additionally, while the lac-reversion assay could be an alternative, it focuses on specific metabolic pathway mutations (such as those affecting lacZ) and would not necessarily capture the same types of mutations relevant to rifampicin resistance or antibiotic-induced mutagenesis. Given our experimental objective of understanding how ampicillin induces mutations that confer antibiotic resistance, the current approach of using rifampicin selection provides a direct and relevant measurement of mutation frequency under antibiotic stress.

b. Major: It is important to establish what the basal mutation frequencies/rates of both the WT and ∆recA strains are. Currently, only the ampicillin-treated populations were reported. It is possible that the ∆recA strain has an inherently higher mutagenesis than WT. Thus, ampicillin treatment might not in fact induce higher mutagenesis in ∆recA.

Thanks for your suggestion. The basal mutation frequency of the wild-type and the ∆recA strain have been measured using rifampicin (Fig. 1G), and there is no significant difference between them.

c. Major: In the text, the authors write, "To verify whether drug resistance associated DNA mutations have led to the rapid development of antibiotic resistance in recA mutant strain, we randomly selected 15 colonies on non-selected LB agar plates from the wild type surviving isolates, and antibiotic screening plates containing 50 μg/mL ampicillin from the ΔrecA resistant isolates, respectively." Why were the WT clones picked from non-selective plates and the recA mutant from selective ones for WGS? It appears that such a procedure would bias the recA mutant clones to show more mutations (caused by selection on the ampicillin plate). The authors need to address this discrepancy.

We appreciate your concern regarding potential inconsistencies in the WGS methodology. However, we would like to clarify that the primary aim of the WGS experiment was to identify the types of mutations present in the wild-type and ΔrecA strains after treatment of ampicillin, rather than to quantify or compare mutation frequencies. This purpose was explicitly stated in the manuscript.

Furthermore, the choice of selective and non-selective conditions was made to ensure the successful isolation of mutants in both strains. Specifically, if selective conditions (50 μg/mL ampicillin) were applied to the wild type strain, it would have been nearly impossible to recover colonies for WGS analysis, as wild-type cells are highly susceptible to ampicillin at this concentration (Top, Author response image 1). Conversely, under non-selective conditions, ΔrecA mutants carrying resistance mutations may not have been effectively isolated, which would have limited our ability to identify resistance mutations in these strains (Bottom, Author response image 1). Thus, the use of different selection pressures was essential for achieving the objective of mutation identification in this study.

d. Major: In some instances, the authors do not use accurate language to describe their data. In Figure 2A, the authors randomly selected 15 ∆recA clones from a selective plate with 50 µg/mL of ampicillin. These clones were then subjected to WGS, which subsequently identified resistant mutations. Based on the described methods, these mutations are a result of selection: in other words, resistant mutations were preexisting in the bacterial population, and the addition of ampicillin selection killed off the sensitive cells, enabling the proliferation of the resistant clones. However, the in Figure 2 legend and associated text, the authors suggest that these mutations were "induced" by beta-lactam exposure, which is misleading. The data does not support that.

We appreciate your detailed feedback on the language used to describe our data. We understand the concern regarding the use of the term "induced" in relation to beta-lactam exposure. To clarify, we employed not only beta-lactam antibiotics but also other antibiotics, such as ciprofloxacin and chloramphenicol, in our experiments (data not shown). However, we observed that beta-lactam antibiotics specifically induced the emergence of resistance or altered the MIC in our bacterial populations. If resistance had pre-existed before antibiotic exposure, we would expect other antibiotics to exhibit a similar selective effect, particularly given the potential for cross-resistance to multiple antibiotics.

Furthermore, we used two different ∆recA strains, and the results were consistent between the strains (Fig. S3). Given that spontaneous mutations can occur with significant variability in populations, if resistance mutations pre-existed before antibiotic exposure, the selective outcomes should have varied between the two strains.

Most importantly, we found that the addition of anti-oxidative compound GSH prevented the evolution of antibiotic from the treatment of ampicillin in the ΔrecA strain. If we assume that resistant bacteria preexist in the ∆recA strain, then the addition of GSH should not affect the evolution of resistance. Therefore, we believe that the resistance mutations we detected were not simply the result of selection from preexisting mutations but were indeed induced by beta-lactam exposure.

e. Major: For Figure 4J, using WGS the authors show that the addition of GSH to WT and ∆recA cells inhibited the rise of resistance mutations; no resistance mutations were reported. However, in the "Whole genome sequencing" section under "Materials and Methods", they state that "Resistant clones were isolated by selection using LB agar plates with the supplementation of ampicillin at 50 μg/mL". These clones were then genome-extracted and sequenced. Given the methodology, it is surprising that the WGS did not reveal any resistance mutations in the GSH-treated cells. How were these cells able to grow on 50 μg/mL ampicillin plates for isolation in the first place? The authors need to address this.

We sincerely apologize for the confusion caused by the incorrect expression in the "Materials and Methods" section. Indeed, when bacteria were treated with the combination of antibiotics and GSH, resistance was significantly suppressed, and no resistant clones could be isolated from selective plates (i.e., LB agar supplemented with 50 μg/mL ampicillin).

To address this, we instead plated the bacteria treated with antibiotics and GSH onto non-selective plates (without ampicillin) and randomly selected 15 colonies for WGS. None of them showed resistance mutations. We will revise the text in the "Materials and Methods" section to accurately reflect this procedure and provide clarity.

f. Minor: for Figure 1G, it is misleading to have both "mutation frequency" and "mutant rate" in the y-axis; the two are defined and calculated differently. Based on the Materials and Materials, "mutation frequency" would be the appropriate term. Also, for the ∆recA strain, it is a bit unusual to see mutation frequencies that are tightly clustered. Usually, mutation frequencies follow the Luria-Delbruck distribution. Can the authors explain why the ∆recA data looks so different compared to, say, the WT mutation frequencies?

Thank you for your insightful feedback. We agree that having both "mutation frequency" and "mutant rate" on the y-axis is misleading, as these terms are defined and calculated differently. To avoid confusion, we will revise Figure 1G to use only "mutation frequency" as the correct term, in line with the methods described in the Materials and Methods section.

Regarding the ∆recA strain's mutation frequencies, we acknowledge that the data appear more tightly clustered compared to the expected Luria-Delbruck distribution seen in the wild type strain. In fact, the y-axis of the Figure 1G is logarithmic, this causes the data to appear more clustered.

We further added the basal mutation frequency in the wild type and ∆recA strains before the exposure to ampicillin. The basal mutation frequency of the wild-type and the ∆recA strain have been measured using rifampicin (Fig. 1G), and there is no significant difference between them.

g. Minor: It needs to be made clear in the Main Text what the selective antibiotic agar plate used was, rifampicin or ampicillin. I am assuming it was rifampicin, as ampicillin plates would yield resistance frequencies close to 100%, given the prior treatment of the culture with ampicillin.

Thanks for your comments. Depending on the objective, we used different selective plates. For example, when testing the mutation frequency of antibiotic resistance, we used a selective plate containing rifampicin in order to utilize a non-inducing antibiotic, which is the standard method for calculating resistance mutation frequency. In the WGS experiment, to obtain mutations specific to ampicillin resistance, we selected a selective plate containing ampicillin.

Reviewer #2:

The Y-axis label (log10 mutant rate) in Figure 1G is misleading or incorrect.

Thanks for your comments and we apologize for this misleading information. The Figure 1G has been revised accordingly.

In line 393 of the discussion, the authors claim that excessive ROS accumulation drives the evolution of ampicillin resistance, which has not been conclusively demonstrated. Additional experiments are needed to support this statement.

We greatly appreciate your comments. First, the correlation between DNA mutations and the accumulation of reactive oxygen species (ROS) has been experimentally confirmed. As shown in Fig. 4I, after the addition of the antioxidant GSH, DNA resistance mutations were not detected in the ΔrecA strain treated with ampicillin for 8 hours, compared to those without the addition of GSH, proving that the rapid accumulation of ROS induces the enhancement of DNA resistance mutations. Second, the enhancement of DNA resistance mutations in relation to bacterial resistance has been widely validated and is generally accepted. Finally, we appreciate the your suggestion to strengthen the evidence supporting ROS enhancement. To address this, we have added an experiment to measure ROS levels. Through flow cytometry, we found that ROS levels significantly increased in both the wild-type and ΔrecA strain after 8 hours of ampicillin treatment. However, ROS levels in the ΔrecA strain showed a significant further increase compared to the wild-type strain (Fig. 4G). Additionally, with the addition of 50 mM glutathione, no significant change in ROS levels was observed in either the wild-type or ΔrecA strain before and after ampicillin treatment (Fig. 4H). This result further confirms our finding in Fig. 4I, where adding GSH inhibited the development of antibiotic resistance.

The abstract is overly complex and difficult to read, e.g. "Contrary to previous findings, it is shown that this accelerated resistance development process is dependent on the hindrance of DNA repair, which is completely orthogonal to the SOS response").

Thank you for the valuable feedback regarding the complexity of the abstract. We agree that certain sections could be simplified for clarity. In response, we have revised the abstract to make it more concise and easier to understand. For example, the sentence “Contrary to previous findings, it is shown that this accelerated resistance development process is dependent on the hindrance of DNA repair, which is completely orthogonal to the SOS response” has been rewritten as: "Unlike earlier studies, we found that the rapid development of resistance relies on the hindrance of DNA repair, a mechanism that operates independently of the SOS response."

Reviewer #3:

As indicated above, direct evidence is needed to show (1) that these phenotypes exist in strains harboring deletions in other DNA repair genes outside of the SOS response, (2) that DNA damage is increased, (3) that reactive oxygen species accumulate, (4) that accelerated resistance evolution can be reversed by anything other than recA complementation. There are also other resistance evolution mechanisms untested here, including transcription-coupled repair (TCR) mechanisms involving Mfd. These need to be shown in order to draw the conclusions proposed.

We sincerely thank you for your insightful comments. First, in this study, our primary focus is on the role of recA deficiency in bacterial antibiotic resistance evolution. Therefore, we conducted an in-depth investigation on E. coli strains lacking RecA and found that its absence promotes resistance evolution through mechanisms involving increased ROS accumulation and downregulation of DNA repair pathways. While we acknowledge the importance of other DNA repair genes outside of the SOS response and other resistance evolution mechanisms including the TCR mechanism, exploring them is beyond the scope of this paper. However, in a separate unpublished study, we have identified the involvement of another DNA recombination protein, whose role in resistance evolution is not yet fully elucidated, in promoting resistance development. This finding is part of another independent investigation.

Regarding DNA damage and repair, our paper emphasizes that resistance-related mutations in DNA are central to the development of antibiotic resistance. These mutations are a manifestation of DNA damage. To demonstrate this, we measured mutation frequency and performed whole-genome sequencing, both of which confirmed an increase in DNA mutations.

We appreciate the reviewer's suggestion to provide additional evidence for ROS accumulation, and we have now supplemented our manuscript with relevant experiments. Through flow cytometry, we found that ROS levels significantly increased in both the wild type and ΔrecA strains after 8 hours of ampicillin treatment. However, ROS levels in the ΔrecA strain showed a significant further increase compared to the wild-type strain (Fig. 4G). Additionally, with the addition of 50 mM glutathione, no significant change in ROS levels was observed in either the wild-type or ΔrecA strain before and after ampicillin treatment (Fig. 4H). This result further confirms our finding in Fig. 4I, where adding GSH inhibited the development of antibiotic resistance.

Finally, in response to your question about reversing accelerated resistance evolution, we would like to highlight that, in addition to recA complementation, we successfully suppressed rapid resistance evolution by supplementing with an antioxidant, GSH (Fig. 4I). This further supports our hypothesis that increased ROS levels play a key role in driving accelerated resistance evolution in the absence of RecA.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The killing mechanism of many antibiotics involves the induction of DNA damage, either directly or indirectly, which triggers the SOS response. RecA, the master regulator of the SOS response, plays a crucial role in driving the evolution of resistance to fluoroquinolone antibiotics treated with a single dose of ciprofloxacin. However, the precise roles of RecA and SOS responses in the development of resistance under short-term β-lactam exposure remain unclear. In the present study, we observed a fast evolution of β-lactam resistance (20-fold increase in MIC in 8 hours) in E. coli after deleting RecA and exposing the bacteria to a single dose of ampicillin. Notably, once this type of resistance is established, it remains stable and can be passed on to subsequent generations. Contrary to previous findings, it is shown that this accelerated resistance development process is dependent on the hindrance of DNA repair, which is completely orthogonal to the SOS response. Additionally, we identified the rapid emergence of drug resistance associated mutations in the resistant bacterial genome, indicating the impairment of DNA repair. Through comprehensive transcriptome sequencing, we discovered that the expression of numerous antioxidative response genes is repressed in recA mutant resistant isolates, resulting in an excessive accumulation of ROS within the cells. This suggests that the induction of ROS drives the fast evolution of antibiotic resistance in RecA-deficient bacteria. Collectively, we show that the hindrance of DNA repair hampers cellular fitness, provides bacteria with genetic adaptability to survive in diverse stressful environments, and accelerates the evolution of antibiotic resistance.

Revised version: The killing mechanism of many antibiotics involves the induction of DNA damage, either directly or indirectly, which triggers the SOS response. RecA, the master regulator of the SOS response, plays a crucial role in driving the evolution of resistance to fluoroquinolone antibiotics treated with a single dose of ciprofloxacin. However, the precise roles of RecA and SOS responses in the development of resistance under short-term β-lactam exposure remain unclear. In the present study, we observed a fast evolution of β-lactam resistance (20-fold increase in MIC in 8 hours) in E. coli after deleting RecA and exposing the bacteria to a single dose of ampicillin. Notably, once this type of resistance is established, it remains stable and can be passed on to subsequent generations. Unlike earlier studies, we found that the rapid development of resistance relies on the hindrance of DNA repair, a mechanism that operates independently of the SOS response. Additionally, we identified the rapid emergence of drug resistance associated mutations in the resistant bacterial genome, indicating the impairment of DNA repair. Through comprehensive transcriptome sequencing, we discovered that the expression of numerous antioxidative response genes is repressed in recA mutant resistant isolates, resulting in an excessive accumulation of ROS within the cells. This suggests that the induction of ROS drives the fast evolution of antibiotic resistance in RecA-deficient bacteria. Collectively, we show that the hindrance of DNA repair hampers cellular fitness, provides bacteria with genetic adaptability to survive in diverse stressful environments, and accelerates the evolution of antibiotic resistance.

Introduction

Addressing bacterial infections caused by emerging and drug-resistant pathogens represents a major global health priority. Bactericidal antibiotics can exert their effects on cells by directly or indirectly causing DNA damage or triggering the production of highly destructive hydroxyl radicals (1–3). This, in turn, initiates a protective mechanism known as the SOS response, which enables bacterial survival against the lethal impacts of antibiotics by activating intrinsic pathways for DNA repair (4–7). The activation of DNA repair processes relies on specific genes, such as recA, which encodes a recombinase involved in DNA repair, and lexA, a repressor of the SOS response that can be inactivated by RecA (8).

Studies have demonstrated that a single exposure to fluoroquinolones, a type of antibiotic that induces DNA breaks and triggers the SOS response, leads to the development of bacterial resistance in Escherichia coli (E. coli) through a RecA and SOS response-dependent mechanism (9). Given the crucial role of RecA in the SOS response, inhibiting RecA activity to deactivate the SOS response presents an appealing strategy for preventing the evolution of bacterial resistance to antibiotics (10). Similarly, exposure to fluoroquinolones induces the SOS response and mutagenesis in Pseudomonas aeruginosa, and the deletion of recA in this pathogen results in a significant reduction in resistance to fluoroquinolones (11).

Unlike fluoroquinolones, β-lactam antibiotics induce a RecA-dependent SOS response in E. coli through impaired cell wall synthesis, mediated by the DpiBA two-component signal system (12). The development of antibiotic resistance, triggered by exposure to β-lactams, has been extensively investigated using the cyclic adaptive laboratory evolution (ALE) method. Mutations that arise during cyclic ALE experiments are attributed to errors occurring during continued growth, necessitating multiple rounds of β-lactam exposure to drive the evolution of resistance in E. coli cells (13,14). However, the precise roles of RecA and SOS responses in the development of resistance under short-term β-lactam antibiotics exposure remain unclear.

Recently, there has been a growing interest in understanding the impact of the stress-induced accumulation of reactive oxygen species (ROS) on bacterial cells (15,16). While exploring methods to harness ROS-mediated killing has the potential to enhance the effectiveness of various antibiotics (17–19), the role of ROS in antimicrobial activity has become a topic of controversy following challenges to the initial observations (20,21). The generation of ROS has been found to contribute to the development of multidrug resistance, as prolonged exposure to antibiotics in cyclic ALE experiments is known to generate ROS, leading to DNA damage and increased mutagenesis (22,23). Nevertheless, there is still limited knowledge regarding the consequences of ROS accumulation in bacteria when the activity of RecA or the SOS response is suppressed. Here, in the present work, we observed a fast evolution of multi-drug resistance in an E. coli strain with the recA mutation, following a single exposure to β-lactam antibiotics. This evolution can be attributed to three mechanisms: compromised DNA repair, transcriptional repression of genes associated with antioxidative processes, and the excessive accumulation of ROS.

Results

Single exposures to β-lactam antibiotics trigger a fast evolution of antibiotic resistance in the recA mutant strain

Earlier version: To investigate the impact of the SOS response on bacterial evolution towards β-lactam resistance, we generated a recA mutant strain (ΔrecA) from the E. coli MG1655 strain (LZ101). Initially, we conducted an ALE experiment using a slightly modified treatment protocol (24) on the wild type and ΔrecA strains. During a period of three weeks, the cells were subjected to cycles of ampicillin exposure for either 4 or 8 hours at a concentration of 50 µg/mL (10 times the MIC) each day (Fig. S1A) (25). As anticipated based on previous studies, the intermittent ampicillin treatments over the course of three weeks resulted in the evolution of antibiotic resistance in the wild type strain (Fig. S1B and C). However, we observed rapid development of ampicillin resistance in the ΔrecA strain after a single exposure for 8 hours to ampicillin (Fig. 1A-C). To ensure that the emergence of resistance we observed was not illusory due to technical issues during the recA knockout process, we employed another ΔrecA strain (JW2669-1) provided by the Coli Genetic Stock Centre (CGSC) with the same killing procedure. The results from both bacterial strains were consistent (Fig. S2A and B). To further investigate, we treated both the wild type and ΔrecA cells with other β-lactams, including penicillin G and carbenicillin, at concentrations equivalent to 10 times the MIC (1 mg/mL and 200 µg/mL, respectively) for 8 hours (26,27). Consistently, these treatments also led to a fast evolution of antibiotic resistance in the ΔrecA strain (Fig. S3A and B).

Revised version: To investigate the impact of the SOS response on bacterial evolution towards β-lactam resistance, we generated a recA mutant strain (ΔrecA) from the E. coli MG1655 strain (LZ101). Initially, we conducted an ALE experiment using a slightly modified treatment protocol (24) on the wild type and ΔrecA strains. During a period of three weeks, the cells were subjected to cycles of ampicillin exposure for 4.5 hours at a concentration of 50 µg/mL (10 times the MIC) each day (Fig. S1A) (25). As anticipated based on previous studies, the intermittent ampicillin treatments over the course of three weeks resulted in the evolution of antibiotic resistance in the wild type strain (Fig. S1B). However, a significantly accelerated development of resistance was found in the ΔrecA strain, with the average time to resistance reduced to 2 days (Fig. S1B). More importantly, we observed that resistance even emerged in the ΔrecA strain after a single exposure for 8 hours to ampicillin (Fig. 1A-C). Meanwhile, after 8 hours of treatment with 50 μg/mL ampicillin, the survival rates of both wild type and ΔrecA strain were consistent (Fig. S2). To ensure that the emergence of resistance we observed was not illusory due to technical issues during the recA knockout process, we employed another ΔrecA strain (JW2669-1) provided by the Coli Genetic Stock Centre (CGSC) with the same killing procedure. The results from both bacterial strains were consistent (Fig. S3A and B). To further investigate, we treated both the wild type and ΔrecA cells with other β-lactams, including penicillin G and carbenicillin, at concentrations equivalent to 10 times the MIC (1 mg/mL and 200 µg/mL, respectively) for 8 hours (26,27). Consistently, these treatments also led to a fast evolution of antibiotic resistance in the ΔrecA strain (Fig. S4A and B).

Fast evolution of antibiotic resistance in E. coli recA mutant strain.

Earlier version: (A) Experimental flow for the single exposures to antibiotics in E. coli strains. Step i: an overnight culture (1 × 109 CFU/mL cells) was diluted 1:50 into 30 mL LB medium supplemented with 50 μg/mL ampicillin and incubated at 37°C with shaking at 250 rpm for 0, 2, 4, 6 and 8 hours; Step ii: after each treatment, the ampicillin containing medium was removed by washing twice in a fresh LB medium; Step iii: the surviving isolates were resuspended in 30 mL fresh LB medium and regrown overnight at 37°C with shaking at 250 rpm; Step iv: cell cultures were plated onto LB agar and incubated for 16 hours at 37°C; Step v: single colonies were inoculated in 30 mL fresh LB medium and cultured at 37°C with shaking at 250 rpm for 4 to 6 hours. (B) MICs of ampicillin were measured against the wild type E coli strain after single exposures to ampicillin. (C) MICs of ampicillin were measured against the ΔrecA strain after single exposures to ampicillin. (D) After the treatment of Step v, cells were continuously cultured in an antibiotic-free medium for seven days. MICs of ampicillin were measured each day. (E) MICs of ampicillin were measured against the ΔrecA strain treated with ampicillin, where the expression of RecA was restored using plasmid-based constitutive expression of recA before the treatment of Step i. (F) MICs of rifampicin were measured after an 8-hour treatment of ampicillin for the wild type and ΔrecA strain. (G) Mutation frequencies using rifampicin as the selector were measured and calculated in the wild type and ΔrecA strain after the single exposure to ampicillin. Each experiment was independently repeated at least six times using parallel replicates, and the data are shown as mean ± SEM. Significant differences among different treatment groups are analyzed by independent t-test, *P < 0.05, **P < 0.01, ***P < 0.001.

Revised version: (A) Experimental flow for the single exposures to antibiotics in E. coli strains. Step i: an overnight culture (1 x 109 CFU/mL cells) was diluted 1:50 into 30 mL LB medium supplemented with 50 μg/mL ampicillin and incubated at 37°C with shaking at 250 rpm for 0, 2, 4, 6 and 8 hours; Step ii: after each treatment, the ampicillin containing medium was removed by washing twice in a fresh LB medium; Step iii: the surviving isolates were resuspended in 30 mL fresh LB medium and regrown overnight at 37°C with shaking at 250 rpm; Step iv: cell cultures were plated onto LB agar and incubated for 16 hours at 37°C; Step v: single colonies were inoculated in 30 mL fresh LB medium and cultured at 37°C with shaking at 250 rpm for 4 to 6 hours. (B) MICs of ampicillin were measured against the wild type E coli strain after single exposures to ampicillin. (C) MICs of ampicillin were measured against the ΔrecA strain after single exposures to ampicillin. (D) After the treatment of Step v, cells were continuously cultured in an antibiotic-free medium for seven days. MICs of ampicillin were measured each day. (E) MICs of ampicillin were measured against the ΔrecA strain treated with ampicillin, where the expression of RecA was restored using plasmid-based constitutive expression of recA before the treatment of Step i. (F) MICs of rifampicin were measured after an 8-hour treatment of ampicillin for the wild type and ΔrecA strain. (G) Mutation frequencies using rifampicin as the selector were measured and calculated in the wild type and ΔrecA strain before and after the single exposure to ampicillin. Each experiment was independently repeated at least six times using parallel replicates, and the data are shown as mean ± SEM. Significant differences among different treatment groups are analyzed by independent t-test, *P < 0.05, **P < 0.01, ***P < 0.001.

To assess the stability of this accelerated antibiotic resistance acquired by the ΔrecA strain, we conducted a study wherein the ΔrecA resistant isolates, originating from the initial 8-hour treatment with ampicillin, were continuously cultivated in a medium devoid of antibiotics for a period of seven days. Our findings revealed that once resistance was established, resistance remained stable and was able to be passed on to subsequent generations even in the absence of ampicillin (Fig. 1D). Moreover, we performed a complementation experiment by introducing a plasmid containing recA under its native promoter into the ΔrecA strain prior to Step i in Fig. 1A, that is, before exposing the cells to ampicillin. Interestingly, this complemented strain displayed a comparable MIC to the isogenic wild type strain and maintained its sensitivity even after ampicillin treatment for up to 8 hours (Fig. 1E).

Earlier version: Notably, the evolution of antibiotics is influenced by mutation and selection, among other factors. To further investigate the primary drivers behind the rapid development of antibiotic resistance, we employed rifamycin, an unrelated antibiotic. Our findings indicated that following an 8-hour exposure to ampicillin, the MIC of rifampicin reached about 100 µg/mL for the ΔrecA strain, marking a 10-fold increase compared to the wild-type strain (Fig. 1F). Moreover, we measured and calculated the mutation frequency by employing rifampicin as the selective agent in both the wild type and ΔrecA strains after the 8-hour ampicillin exposure. This was done by dividing the number of colony-forming units (CFU) per milliliter on ampicillin plates by the number of CFU per milliliter on LB plates (28). Our results revealed a significantly higher mutation frequency in the ΔrecA strain in contrast to the wild type (Fig. 1G). These findings together demonstrate that a singular exposure to β-lactam antibiotics markedly boosts mutation frequency and thereby triggers a fast evolution of antibiotic resistance that is stable and heritable in the ΔrecA strain.

Revised version: Notably, the evolution of antibiotics is influenced by mutation and selection, among other factors. To further investigate the primary drivers behind the rapid development of antibiotic resistance, we employed rifamycin, an unrelated antibiotic. Our findings indicated that following an 8-hour exposure to ampicillin, the MIC of rifampicin reached about 100 µg/mL for the ΔrecA strain, marking a 10-fold increase compared to the wild-type strain (Fig. 1F). Moreover, we measured and calculated the mutation frequency by employing rifampicin as the selective agent in both the wild type and ΔrecA strains after the 8-hour ampicillin exposure. This was done by dividing the number of colony-forming units (CFU) per milliliter on rifampicin plates by the number of CFU per milliliter on LB plates (28). Our results revealed a significantly higher mutation frequency in the ΔrecA strain in contrast to the wild type (Fig. 1G). These findings together demonstrate that a singular exposure to β-lactam antibiotics markedly boosts mutation frequency and thereby triggers a fast evolution of antibiotic resistance that is stable and heritable in the ΔrecA strain.

Rapid appearance of drug resistance associated DNA mutations in recA mutant resistant isolates

In bacteria, resistance to most antibiotics requires the accumulation of drug resistance associated DNA mutations developed over time to provide high levels of resistance (29). To verify whether drug resistance associated DNA mutations have led to the rapid development of antibiotic resistance in recA mutant strain, we randomly selected 15 colonies on non-selected LB agar plates from the wild type surviving isolates, and antibiotic screening plates containing 50 µg/mL ampicillin from the ΔrecA resistant isolates, respectively, and performed whole-genome sequencing. We found that abundant drug resistance associated mutations occurred within all resistant isolates, including the promoter of the β-lactamase ampC (PampC) in 8 isolates, the ampicillin-binding target PBP3 (ftsI) in 1 isolate, and the AcrAB-TolC subunit AcrB (acrB) mutations in 6 isolates (Fig. 2A). A mutation in gene nudG was detected in wild type surviving isolates after the single exposure to ampicillin (Fig. 2A), which is involved in pyrimidine (d)NTP hydrolysis to avoid DNA damage (30). Other mutations were listed in the Table S1.

Rapid induction of drug resistance associated DNA mutations in the recA mutant strain.

(A) Detection of drug resistance associated DNA mutations in the wild type and ΔrecA strain after the single exposures to ampicillin at 50 µg/mL for 8 hours. (B) MICs of chloramphenicol against the wild type and ΔrecA strain after the single treatment with ampicillin were tested. (C) MICs of kanamycin against the wild type and ΔrecA strain after the single treatment with ampicillin were tested. (D) The wild type, ΔrecA, and ΔrecAacrB mutant resistant isolates were incubated with NMP at different concentrations for 12 hours. Subsequently, MICs were tested in these strains for resistance to ampicillin. Each experiment was independently repeated at least six times, and the data is shown as mean ± SEM. Significant differences among different treatment groups are analyzed by independent t-test, *P < 0.05, **P < 0.01, ***P < 0.001, ns, no significance.

Earlier version: The presence of PampC mutations was accompanied by a significant increase in the production of β-lactamase in bacteria (Fig. S4). This leads to specific resistance to β-lactam antibiotics. The gene acrB codes for a sub-component of the AcrAB-TolC multi-drug efflux pump, which is central in Gram-negative bacteria (28,31). Mutations in AcrAB-TolC enhance the efflux of antibiotics and confer resistance to multiple drugs (28). Consequently, after short-term exposure to ampicillin, the ΔrecA isolates carrying the acrB mutations exhibited resistance to other types of antibiotics, such as chloramphenicol and kanamycin (Fig. 2B and C). Treatment with high concentrations of 1-(1-Naphthylmethyl) piperazine (NMP), an efflux pump inhibitor (EPI) that competitively blocks TolC-composed efflux pumps, successfully restored the sensitivity of ΔrecA resistant isolates to ampicillin, bringing them back to the equivalent concentrations found in the wild type (Fig. 2D). Collectively, these results suggest that a single exposure to β-lactam antibiotics can lead to drug resistance associated DNA mutations that play a crucial role in the accelerated development of antibiotic resistance in the ΔrecA strain.

Revised version: The presence of PampC mutations was accompanied by a significant increase in the production of β-lactamase in bacteria (Fig. S5). This leads to specific resistance to β-lactam antibiotics. The gene acrB codes for a sub-component of the AcrAB-TolC multi-drug efflux pump, which is central in Gram-negative bacteria (28,31). Mutations in AcrAB-TolC enhance the efflux of antibiotics and confer resistance to multiple drugs (28). Consequently, after short-term exposure to ampicillin, the ΔrecA isolates carrying the acrB mutations exhibited resistance to other types of antibiotics, such as chloramphenicol and kanamycin (Fig. 2B and C). Treatment with high concentrations of 1-(1-Naphthylmethyl) piperazine (NMP), an efflux pump inhibitor (EPI) that competitively blocks TolC-composed efflux pumps, successfully restored the sensitivity of ΔrecA resistant isolates to ampicillin, bringing them back to the equivalent concentrations found in the wild type (Fig. 2D). Collectively, these results suggest that a single exposure to β-lactam antibiotics can lead to drug resistance associated DNA mutations that play a crucial role in the accelerated development of antibiotic resistance in the ΔrecA strain.

Hindrance of SOS-independent DNA repair in recA mutant resistant isolates

The impaired ability to repair DNA damage can lead to an accelerated accumulation of DNA mutations. The cellular response to DNA damage, such as the SOS response, often influences the balance between bacterial susceptibility and resistance. We thereafter asked whether the absence of recA and subsequently induced inhibition of the SOS response can contribute to the rapid development of antibiotic resistance.

To investigate this, we first tested the ability of various mutants involved in different pathways of the SOS response to evolving antibiotic resistance following a single treatment with ampicillin for 8 hours at 50 µg/mL. A mutant form of the SOS master regulator LexA (lexA3), which is incapable of being cleaved and thus defective in SOS induction, did not exhibit antibiotic resistance evolution (Fig. 3A). Additionally, the deletion of either DpiB or DpiA (encoded by citB or citA, respectively), inhibiting the DpiBA two-component signal transduction system, did not result in the development of antibiotic resistance after ampicillin exposure (Fig. 3A). Moreover, the deletion of several downstream effectors of the SOS response, including those involved in cell division inhibition (SulA and YmfM encoded by sulA and ymfM) (32) (Fig. 3A), and DNA repair (DNA Pol II, DNA Pol IV, and DNA Pol V encoded by polB, dinB and umuDC) (33) also did not lead to the evolution of antibiotic resistance (Fig. 3A). Interestingly, these findings indicate that the fast evolution of antibiotic resistance is supposed to occur in an SOS-independent manner in the absence of recA.

SOS-independent impairment of DNA repair in ΔrecA resistant isolates.

(A) MICs of ampicillin were tested against the wild type strain, ΔrecA strain, and mutants lacking specific genes from the SOS response after single exposures to ampicillin at 50 µg/mL for 8 hours. (B) Filament cell lengths in the wild type (n=253) and the ΔrecA strain (n=216) after single treatments with ampicillin at 50 μg/mL. (C) Multinucleated filaments were observed in the wild type (i) and the ΔrecA (ii) strain after single exposures to ampicillin at 50 μg/mL. Purple: E. coli chromosome; green: DNA Pol I. (D) Relative fold changes of DNA Pol I in the wild type and ΔrecA strain after single treatments with ampicillin at 50 μg/mL. (E) Co-localization between the E. coli chromosome and DNA Pol I in the wild type and ΔrecA strain after the single exposures to ampicillin at 50 μg/mL. Data is shown as mean ± SEM. Significant differences among different treatment groups are analyzed by independent t-test, *P < 0.05, **P < 0.01, ***P < 0.001, ns, no significance.

In addition to the DNA repair components associated with the SOS response, DNA Pol I plays a role in processing RNA primers during lagging-strand synthesis and filling small gaps during DNA repair reactions (34). Since DNA Pol I (encoded by polA) has been demonstrated as an essential gene required for the growth of E. coli in rich medium, including the LB medium (35,36), we next utilized Single Molecule Localization Microscopy (SMLM) to precisely locate the chromosome and DNA Pol I in a dynamic manner. During an 8 hour exposure to ampicillin, we observed the formation of multinucleated filaments in both the wild type and ΔrecA strain, indicating a pause in cell division and suggesting a time period for bacterial DNA repair to take place (Fig. 3B and C) (37). However, the expression level of DNA Pol I was significantly suppressed in the ΔrecA strain compared to the wild type strain after 4 hours of ampicillin exposure (Fig. 3D). More notably, the super-resolution colocalization analysis revealed a significantly lower ratio of co-localization between the chromosome and DNA Pol I in the ΔrecA strain (Fig. 3E), which together demonstrate that the hindrance of DNA repair caused by the absence of recA is crucial for the fast evolution of antibiotic resistance in an SOS-independent manner.

Repression on antioxidative-related gene transcription drives the fast evolution of antibiotic resistance in the recA mutant strain

To further comprehend the fast evolution of β-lactam resistance that was observed in this study, we investigated the changes in gene transcription induced by ampicillin treatment using a comprehensive transcriptome sequencing approach (total RNA-seq). Our analysis revealed significant alterations in the transcriptomic profiles of both the wild type and ΔrecA strain isolates following a single treatment with ampicillin, compared to untreated controls (Fig. 4A). Specifically, we identified changes in the expression of 161 and 248 coding sequences (with log2FC > 2 and P value < 0.05) in the wild type and ΔrecA strains, respectively. Principal component analysis (PCA) demonstrated a notable disparity in the effects of ampicillin on the ΔrecA strain compared to the wild type strain (Fig. 4B). Additionally, a Venn diagram analysis confirmed that 138 and 225 genes were uniquely regulated by ampicillin exposure in the wild type and ΔrecA strains, respectively (Fig. 4C).

Overaccumulation of ROS drives the fast evolution of multi-drug resistance in the ΔrecA strain.

Earlier version: (A) Clustered heatmap of relative expression of coding sequences in the wild type and ΔrecA strain with significant fold changes (log2FC > 2 and P value < 0.05). (B) Principal-component analysis (PCA) of normalised read counts for all strains. (C) Venn diagram of differentially expressed genes (log2FC > 2) after treatment with ampicillin at 50 μg/ml for 8 hours in the wild type and ΔrecA strain. (D) The top 10 most differentially expressed genes in the ΔrecA strain after the single treatment with ampicillin are labelled in each plot. Blue dots indicate genes with a significant downregulation compared to the untreated control (log2FC > 2 and P value < 0.05), and yellow dots indicate genes with a significant upregulation compared to the untreated control (log2FC > 2 and P value < 0.05). (E) Levels of transcription of SOS response system-associated genes and gene polA in the wild type and ΔrecA strain after single exposures to ampicillin for 8 hours. (F) Levels of transcription of different antioxidative associated genes in the wild type and ΔrecA strain after single exposures to ampicillin for 8 hours. (G) The addition of 50 mM antioxidative compound GSH prevented the evolution of antibiotic resistance to ampicillin in the ΔrecA strain treated with ampicillin at 50 μg/ml for 8 hours. (H) Survival fraction after a single exposure to ampicillin at 50 μg/ml for 8 hours in the wild type and the ΔrecA strain with or without the addition of GSH at 50 mM. (I) Levels of transcription of proteins involved in the BER DNA repair system in the wild type and ΔrecA strain after single exposures to ampicillin for 8 hours. (J) Whole genome sequencing confirms undetectable DNA mutations in the wild type and ΔrecA strain treated with single exposures to ampicillin with the addition of GSH at 50 mM for 8 hours. (K) Transcription levels of all transcriptional repressors in the wild type and ΔrecA strain after single treatments with ampicillin for 8 hours. Total RNA-seq was performed with three repeats in each group. Each experiment was independently repeated at least six times, and the data are shown as mean ± SEM. Significant differences among different treatment groups are analyzed by independent t-test, *P < 0.05, **P < 0.01, ***P < 0.001, #P < 0.05.

Revised version: (A) Clustered heatmap of relative expression of coding sequences in the wild type and ΔrecA strain with significant fold changes (log2FC > 2 and P value < 0.05). (B) Principal-component analysis (PCA) of normalised read counts for all strains. (C) Venn diagram of differentially expressed genes (log2FC > 2) after treatment with ampicillin at 50 μg/ml for 8 hours in the wild type and ΔrecA strain. (D) The top 10 most differentially expressed genes in the ΔrecA strain after the single treatment with ampicillin are labelled in each plot. Blue dots indicate genes with a significant downregulation compared to the untreated control (log2FC > 2 and P value < 0.05), and yellow dots indicate genes with a significant upregulation compared to the untreated control (log2FC > 2 and P value < 0.05). (E) Levels of transcription of SOS response system-associated genes and gene polA in the wild type and ΔrecA strain after single exposures to ampicillin for 8 hours. (F) Levels of transcription of different antioxidative associated genes in the wild type and ΔrecA strain after single exposures to ampicillin for 8 hours. (G) ROS levels were measured by flow cytometry before and after 8 hours of ampicillin treatment (50 μg/mL) in the wild type and ΔrecA strains. (H) ROS levels were measured by flow cytometry in the wild type and ΔrecA strains before and after 8 hours of ampicillin treatment at 50 μg/mL with the addition of GSH (50 mM). (I) The addition of 50 mM antioxidative compound GSH prevented the evolution of antibiotic resistance to ampicillin in the ΔrecA strain treated with ampicillin at 50 μg/ml for 8 hours. (J) Survival fraction after a single exposure to ampicillin at 50 μg/ml for 8 hours in the wild type and the ΔrecA strain with or without the addition of GSH at 50 mM. (K) Levels of transcription of proteins involved in the BER DNA repair system in the wild type and ΔrecA strain after single exposures to ampicillin for 8 hours. (L) Whole genome sequencing confirms undetectable DNA mutations in the wild type and ΔrecA strain treated with single exposures to ampicillin with the addition of GSH at 50 mM for 8 hours. (M) Transcription levels of all transcriptional repressors in the wild type and ΔrecA strain after single treatments with ampicillin for 8 hours. Total RNA-seq was performed with three repeats in each group. Each experiment was independently repeated at least six times, and the data are shown as mean ± SEM. Significant differences among different treatment groups are analyzed by independent t-test, *P < 0.05, **P < 0.01, ***P < 0.001, #P < 0.05.

Earlier version: To elucidate the differential expression of genes associated with specific biological functions, we conducted Gene Ontology (GO) enrichment analyses (Fig. S5A and B) and Kyoto Encyclopedia of Genes and Genomes (KEGG) pathway analyses (Fig. S5C and D). Our findings indicate that ampicillin profoundly impacted persistence pathways in the wild type strain, specifically affecting pathways related to quorum sensing, flagellar assembly, biofilm formation, and bacterial chemotaxis (38,39). Conversely, in the ΔrecA strain, a distinct functional category associated with the oxidative stress response exhibited significant and unique down-regulation. This category included activities such as sulfate transporter activity, iron-sulfur cluster assembly, oxidoreductase activity, and carboxylate reductase activity.

Revised version: To elucidate the differential expression of genes associated with specific biological functions, we conducted Gene Ontology (GO) enrichment analyses (Fig. S6A and B) and Kyoto Encyclopedia of Genes and Genomes (KEGG) pathway analyses (Fig. S6C and D). Our findings indicate that ampicillin profoundly impacted persistence pathways in the wild type strain, specifically affecting pathways related to quorum sensing, flagellar assembly, biofilm formation, and bacterial chemotaxis (38,39). Conversely, in the ΔrecA strain, a distinct functional category associated with the oxidative stress response exhibited significant and unique down-regulation. This category included activities such as sulfate transporter activity, iron-sulfur cluster assembly, oxidoreductase activity, and carboxylate reductase activity.

To identify specific genes showing significant fold changes (log2FC > 2 and P value < 0.05), we used volcano plots to visualize the comprehensive changes in gene expression across the genome (Fig. 4D). We examined the transcription levels associated with the SOS response system and found that the transcription of several proteins in the wild type strain can be significantly induced by the single exposure to ampicillin, including citB and dinB (Fig. 4E). However, in the recA mutant strain, antibiotic exposure does not affect the transcription levels of any SOS system-related proteins, suggesting that antibiotic exposure induced the SOS response in the wild type strain but not in the ΔrecA strain. More importantly, we discovered that the induction of the transcription level of the DNA Pol I was significantly suppressed after the single treatment of ampicillin in the ΔrecA strain compared with that in the wild type strain (Fig. 4E). This is consistent with our imaging results and further supports the notion that an SOS-independent evolutionary mechanism dominates the development of antibiotic resistance in recA mutant E. coli.

Earlier version: Further, significant downregulation in the transcription of antioxidative-related genes in the ΔrecA strain was detected, including cysJ, cysI, cysH, soda, and sufD (Fig. 4F). This downregulation suggested an excessive accumulation of reactive oxygen species (ROS) due to compromised cell antioxidative defences. It has been previously reported that the induction of mutagenesis can be stimulated by the overproduction of ROS during antibiotic administration, leading to the evolution of antibiotic resistance both in vivo and in vitro (22,23). Therefore, we hypothesized that elevated ROS levels might be responsible for the observed fast evolution of antibiotic resistance in the ΔrecA strain. To test it, we supplemented the wild type and ΔrecA strains with 50 mM glutathione (GSH), a natural antioxidative compound, and treated them with ampicillin at a concentration of 50 μg/ml for 8 hours. Remarkably, the addition of GSH prevented the development of resistance to ampicillin in the ΔrecA strain (Fig. 4G), without impairing the bactericidal effectiveness of ampicillin (Fig. 4H).

Revised version: Further, significant downregulation in the transcription of antioxidative-related genes in the ΔrecA strain was detected, including cysJ, cysI, cysH, soda, and sufD (Fig. 4F). This downregulation suggested an excessive accumulation of reactive oxygen species (ROS) due to compromised cell antioxidative defences. It has been previously reported that the induction of mutagenesis can be stimulated by the overproduction of ROS during antibiotic administration, leading to the evolution of antibiotic resistance both in vivo and in vitro (22,23). Therefore, we hypothesized that elevated ROS levels might be responsible for the observed fast evolution of antibiotic resistance in the ΔrecA strain. To test it, we first examined the level of ROS generation in the wild type and ΔrecA strains treated with ampicillin for 8 hours by using the fluorescent probe DCFDA/H2DCFDA. We found that ROS levels significantly increased in both the wild-type and ΔrecA strain after 8 hours of ampicillin treatment. However, ROS levels in the ΔrecA strain showed a significant further increase compared to the wild-type strain (Fig. 4G). Additionally, with the addition of 50 mM glutathione (GSH), a natural antioxidative compound, no significant change in ROS levels was observed in either the wild-type or ΔrecA strain before and after ampicillin treatment (Fig. 4H). Further, we supplemented the wild type and ΔrecA strains with 50 mM GSH) and treated them with ampicillin at a concentration of 50 μg/ml for 8 hours. Remarkably, the addition of GSH prevented the development of resistance to ampicillin in the ΔrecA strain (Fig. 4I), without impairing the bactericidal effectiveness of ampicillin (Fig. 4J).

Earlier version: Apart from the SOS response, bacterial cells coordinate other DNA repair activities through a network of regulatory pathways, including excision repair (BER) (40–42). The excessive generation of ROS results in elevated levels of deoxy-8-oxo-guanosine triphosphate (8-oxo-dGTP), an oxidized form of dGTP that becomes both highly toxic and mutagenic upon integration into DNA. The presence of 8-oxo-dG can induce SNP mutations, especially those occurring in guanine, which can be actively rectified by the BER repair pathway (42). Notably, BER glycosylases MutH and MutY can identify and repair these 8-oxo-dG-dependent mutations; however, when MutY and MutH are inactivated, unrepaired 8-oxo-dG can lead to the accumulation of SNP mutations within cells (42). As a result, we conducted a further assessment of the transcription levels of the BER repair pathway in both the wild-type and ΔrecA strain before and after a single exposure to ampicillin. We discovered that after an 8-hour treatment of ampicillin, three DNA repair-associated proteins, including MutH, MutY, and MutM, were notably suppressed in the ΔrecA strain (Fig. 4I).

Revised version: Apart from the SOS response, bacterial cells coordinate other DNA repair activities through a network of regulatory pathways, including base excision repair (BER) (40–42). The excessive generation of ROS results in elevated levels of deoxy-8-oxo-guanosine triphosphate (8-oxo-dGTP), an oxidized form of dGTP that becomes both highly toxic and mutagenic upon integration into DNA. The presence of 8-oxo-dG can induce SNP mutations, especially those occurring in guanine, which can be actively rectified by the BER repair pathway (42). Notably, BER glycosylases MutH and MutY can identify and repair these 8-oxo-dG-dependent mutations; however, when MutY and MutH are inactivated, unrepaired 8-oxo-dG can lead to the accumulation of SNP mutations within cells (42). As a result, we conducted a further assessment of the transcription levels of the BER repair pathway in both the wild-type and ΔrecA strain before and after a single exposure to ampicillin. We discovered that after an 8- hour treatment of ampicillin, three DNA repair-associated proteins, including MutH, MutY, and MutM, were notably suppressed in the ΔrecA strain (Fig. 4K).

Earlier version: Finally, we sequenced the surviving ΔrecA isolates and found that the addition of GSH inhibited drug resistance associated mutations in the ΔrecA strain, which was detected in the ΔrecA resistant isolates including genes of the promoter of ampC and acrB (Fig. 4J). Given the DNA repair impairment resulting in the generation of ROS, it would have been expected for genes involved in the oxidative stress response to be induced in RecA-deficient cells. However, the repression of antioxidative-related genes indicated the involvement of transcriptional repressors that might be regulated by RecA. Consequently, we examined the transcription levels of all transcriptional repressors in both the wild type and ΔrecA strain. Remarkably, we observed a significant upregulation of H-NS, a crucial transcriptional repressor (Fig. 4K). This finding suggests that the upregulation of H-NS could potentially account for the downregulation of antioxidant gene transcription levels in the ΔrecA strain.

Revised version: Finally, we sequenced the surviving ΔrecA isolates and found that the addition of GSH inhibited drug resistance associated mutations in the ΔrecA strain, which was detected in the ΔrecA resistant isolates including genes of the promoter of ampC and acrB (Fig. 4L). Given the DNA repair impairment resulting in the generation of ROS, it would have been expected for genes involved in the oxidative stress response to be induced in RecA-deficient cells. However, the repression of antioxidative-related genes indicated the involvement of transcriptional repressors that might be regulated by RecA. Consequently, we examined the transcription levels of all transcriptional repressors in both the wild type and ΔrecA strain. Remarkably, we observed a significant upregulation of H-NS, a crucial transcriptional repressor (Fig. 4M). This finding suggests that the upregulation of H-NS could potentially account for the downregulation of antioxidant gene transcription levels in the ΔrecA strain.

Discussion

In this study, we present a novel perspective challenging the conventional belief that disabling RecA to inhibit the SOS response can prevent bacteria from developing antibiotic resistance (10). While the SOS response and RecA have been extensively studied for their roles in antibiotic resistance evolution (43), we observed a fast evolution of multi-drug resistance in the ΔrecA strain triggered by single exposures to β-lactam antibiotics, which was driven by three mechanisms, including the impaired DNA repair, transcriptional repression of antioxidative-related genes, and excessive accumulation of ROS (Fig. 5).

Mechanism of rapid development of multi-drug resistance in the recA mutant E. coli strain.

The deletion of recA results in a decrease in DNA damage repair capability. Additionally, the absence of recA upregulates the transcription levels of the transcriptional repressor H-NS, although the mechanism of this regulation is still unclear. Upregulation of H-NS potentially inhibits the transcription levels of several downstream genes associated with antioxidant functions, leading to an excessive accumulation of ROS. The excessive production of ROS causes an increased occurrence of genetic mutations. Under antibiotic pressure, mutations related to drug resistance are selected, ultimately leading to the emergence of rapid multidrug resistance.

We initially suspected that the decrease in DNA repair levels was caused by the suppression of the SOS response, mediated by the deficiency of RecA during rapid resistance development. We therefore examined the response to antibiotic exposure in various E. coli mutants lacking specific genes from the SOS response. However, we discovered that these strains remained susceptible after short-term exposure to ampicillin, providing evidence that the decrease in DNA repair levels, regulated by RecA in an SOS-independent manner, served as a mechanism for the accelerated development of antibiotic resistance. Furthermore, we conducted super-resolution imaging studies to analyze the expression and localization of DNA Pol I. We observed a significant inhibition of DNA Pol I induction and its involvement in DNA repair. Moreover, a notable suppression of the BER repair system was observed in the recA mutant bacteria upon short-term exposure to antibiotics. While we do not yet have evidence explaining why the absence of recA might affect the transcriptional regulation of the BER repair system, this evidence is sufficient to demonstrate that the absence of RecA can lead to the accumulation of beneficial mutations within bacteria through downregulation of DNA repair capacity. The regulation of DNA repair by RecA is a significantly complex process (44), and although we specifically demonstrated the regulation of DNA Pol I and BER repair system by RecA under antibiotic exposure, there should be other repair mechanisms impaired and contributed to this fast evolution of resistance, which needs to be investigated in future.

Antibiotics exert their bactericidal effects by targeting specific intracellular components, and the production of ROS is considered an alternative mechanism underlying quinolone-induced bacterial killing. However, emerging evidence suggests that lethality may arise from primary antibiotic-induced damage or a secondary response to lethal stress mediated by ROS. ROS generation can also contribute to the development of multidrug resistance, as it can directly damage DNA, leading to genetic mutations. In this study, we demonstrate that excessive ROS accumulation drives the evolution of antibiotic resistance in bacteria with impaired DNA repair capabilities. Significantly, we discovered for the first time that antibiotic exposure can significantly increase the transcription of H-NS in the absence of recA. E. coli H-NS is a DNA-binding protein known for its involvement in transcriptional repression (45). It binds to DNA in a sequence-independent manner, with a preference for A-/T-rich regions in curved DNA, thereby impeding transcription (46). Previous findings report that H-NS can down-regulate the transcriptional expression of cysIJH through the mediation of cysB (47,48). While we have established a potential correlation between the increased transcription of H-NS genes induced by antibiotic exposure in RecA deficient cells and the decreased transcription of genes related to antioxidative processes, further investigations are imperative to elucidate the regulatory mechanisms employed by RecA in H-NS transcription and its inhibitory effects on antioxidative gene transcription.

In clinics, the resistance mechanism to β-lactams is not only associated with β-lactamase acquisition but also with increased expression of efflux pumps (44). Our research for the first time demonstrates the fast evolution of antibiotic resistance through the rapid accumulation of acrB mutations and enhanced functionality of the AcrAB-TolC efflux pump in RecA-deficient bacteria exposed to a single dose of ampicillin. This finding supports the clinical observations that the clinical antibiotic resistance is accelerating (50).

DNA repair is a key target for many anti-cancer drugs that aim to destroy tumor cells by inducing DNA damage (51). Homologous recombination (HR) plays a crucial role in DNA repair, and Rad51, a vital protein in the HR pathway of eukaryotic cells, belongs to the recA/RAD51 gene family (51–53). Clinical trials are currently assessing a range of Rad51 inhibitors. Interestingly, studies have revealed that inhibitors targeting Rad51 can effectively inhibit the activity of RecA (54). Despite RecA and Rad51 sharing only approximately 30% sequence homology, the filaments they form and the conformational changes they induce in DNA are remarkably similar (54). Notably, inhibitors of RecA have demonstrated the ability to enhance the antimicrobial effects of certain antibiotics, including fluoroquinolones (55,56). Our findings have significant clinical implications, especially for patients undergoing cancer treatment. The combination therapy involving β-lactam antibiotics and RecA inhibitors could potentially accelerate the evolution of multi-drug resistance at an alarming rate.

Our findings on the mechanism of the fast evolution of multidrug resistance in E. coli after a single exposure to a specific antibiotic can be applied more broadly. The overproduction of ROS is a common effect of many antibiotics and other strategies used to combat infections, such as antimicrobial photodynamic therapy (aPDT) and cold atmospheric plasma (CAP) (57,58). These approaches rely on ROS-mediated damage as the primary source of stress-induced damage. Thus, maintaining the transcription of genes involved in the oxidative stress response or the combined treatment with drugs and antioxidants shows promising potential to prevent stress-induced mutagenesis that is typically triggered by various classes of antibiotics.